```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment03_COMM187S26.ipynb")
```

# COMM 187 S26 Coding Assignment #3

<span style="color:red">**!!! PLEASE MAKE SURE TO RUN THE FIRST CELL ABOVE WHICH SAYS "# Initialize Otter" BEFORE RUNNING AUTO-GRADER !!!**</span>

## <center>TWEETS BY US SENATORS (2008-17)</center>

![](./imgs/A3_introimg.png)

The dataset for this assignment is about the favorites (likes), replies, and retweets on tweets made by US Senators between 2008 and 2017, as recorded on Oct. 19 and 20, 2017. This dataset was collected and analyzed for the following [blog post by FiveThirtyEight](https://fivethirtyeight.com/features/the-worst-tweeter-in-politics-isnt-trump/), and you can read more information about dataset [here](https://github.com/fivethirtyeight/data/tree/master/twitter-ratio/).

The original dataset for this article contains over 288 thousand rows. For the purpose of this assignment, you have beeen provided with a reduced dataset of randomly selected five thousand rows, which is stored in the "data" subfolder and is named `senators_tweets.csv`. It contains the following columns:


Header | Definition
--- | ---
`created_at` | Date and time of the creation of the tweet in the format "mm/dd/yy hh:mm"
`text` | The text content of the tweet
`url` | URL of the tweet
`replies` | Number of replies for the tweet when data was collected
`retweets` | Number of retweets for the tweet when data was collected
`favorites` | Number of favorites (or likes) for the tweet when data was collected
`user` | Twitter username of the US Senator
`bioguide_id` | Unique ID for each US Senator
`party` | Party identification for each US Senator: 'D' for Democrat, 'R' for Republican, and 'I' for Independent
`state` | The state being represented by the US Senator

---

```python
# RUN THIS CELL BEFORE ATTEMPTING THE ASSIGNMENT QUESTIONS
import numpy as np
import pandas as pd
from scipy import stats
```

---

## Question 1: Loading Data (0.5 point)

Use the pandas `read_csv` function to load the data `senator_tweets.csv` file into a pandas DataFrame named `df`.

```python
### Write your code below (in place of ...)
df = ...
```

```python
grader.check("q1")
```

---

## Question 2: Senator Jack Reed's Tweets (1.50 points)

Calculate the following for the tweets made by Senator Jack Reed (user ID = "SenJackReed"):

 - Mean number of retweets, store it in `mean_rt_JackReed` **(0.25 point)**
 - Standard deviation of retweets, using pandas `.std()` method, store it in `std_rt_JackReed` **(0.25 point)**
 - Median number of replies, store it in `median_replies_JackReed` **(0.25 point)**
 - Standard deviation of replies, using numpy `np.std()` function, store it in `std_replies_JackReed` **(0.25 point)**
 - Mean number of favorites, store it in `mean_fav_JackReed` **(0.25 point)**
 - Standard deviation of favorites, using pandas `.std()` method, store it in `std_fav_JackReed` **(0.25 point)**

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q2")
```

---

## Question 3: Comparing favorites across political parties
### A. (0.25 point)

Extract (or subset) `favorites` column for each party (D, R, and I) and store them in respective variables `D`, `R`, and `I`

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q3a")
```

### B. (0.25 point)

Compare the number of **favorites** (or *likes*) between the two major parties (D and R) using t-test. Store the t-statistic in `q3b_t` and the p-value in `q3b_p`.

Steps:
 - Use the variables `D` and `R` defined in the previous question
 - Calculate t-test using `stats.ttest_ind()` function
 - **Tip:** by writing `var1, var2 = stats.ttest_ind()`, you can store the T statistic in `var1` and p value in `var2`

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q3b")
```

### C. (0.25 point)

Compare the number of **favorites** (or *likes*) between the Democratic Senators and Independent Senators (D and I) using t-test. Store the t-statistic in `q3c_t` and the p-value in `q3c_p`.

Steps:
 - Use the variables `D` and `I` defined in the Q3-A.
 - Calculate t-test using `stats.ttest_ind()` function

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q3c")
```

### D. (0.25 point)

Compare the number of **favorites** (or *likes*) between the Republican Senators and Independent Senators (R and I) using t-test. Store the t-statistic in `q3d_t` and the p-value in `q3d_p`.

Follow the same steps as Q3-B and Q3-C. You got this!

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q3d")
```

---

## Question 4: Comparing retweets across states
### A. South Carolina and Arizona (0.5 point)
Compare the **retweets** for the senators from the following two states using t-test: South Carolina (SC) and Arizona (AR).

Store the t statistic in `q4a_t` and the p-value in `q4a_p`. 

Follow similar steps as in Q3-A. You are not being provided with a list of explicit steps for this question, but you have had enough practice to make this work. You got this!


```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q4a")
```

### B. New Jersey and California (0.5 point)
Compare the **retweets** for the senators from the following two states using t-test: New Jersey (NJ) and California (CA).

Store the t statistic in `q4b_t` and the p-value in `q4b_p`.

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q4b")
```

---

Tests for q1 (points: 0.5)
```
>>> all(df == pd.read_csv('./data/senator_tweets.csv'))
True
```
```
>>> type(df) == pd.core.frame.DataFrame
True
```


Tests for q2 (points: 1.5)
```
>>> bool(mean_rt_JackReed == df[df['user'] == 'SenJackReed']['retweets'].mean())
True
```
```
>>> bool(std_rt_JackReed == df[df['user'] == 'SenJackReed']['retweets'].std())
True
```
```
>>> bool(median_replies_JackReed == df[df['user'] == 'SenJackReed']['replies'].median())
True
```
```
>>> bool(std_replies_JackReed == np.std(df[df['user'] == 'SenJackReed']['replies']))
True
```
```
>>> bool(mean_fav_JackReed == df[df['user'] == 'SenJackReed']['favorites'].mean())
True
```
```
>>> bool(std_fav_JackReed == df[df['user'] == 'SenJackReed']['favorites'].std())
True
```


Tests for q3a (points: 0.25)
```
>>> len(D) == 2361
True
```
```
>>> bool(D[19] == 62)
True
```
```
>>> len(R) == 2518
True
```
```
>>> bool(R[1165] == 0)
True
```
```
>>> len(I) == 121
True
```
```
>>> bool(I[3115] == 5018)
True
```


Tests for q3b (points: 0.25)
```
>>> bool(abs(q3b_t - 2.389643559096571) < 0.0001)
True
```
```
>>> bool(abs(q3b_p - 0.016902486599381032) < 0.0001)
True
```


Tests for q3c (points: 0.25)
```
>>> bool(abs(q3c_t + 2.876993743276431) < 0.0001)
True
```
```
>>> bool(abs(q3c_p - 0.0040491404362395355) < 0.0001)
True
```


Tests for q3d (points: 0.25)
```
>>> bool(abs(q3d_t + 8.27417861706419) < 0.0001)
True
```
```
>>> bool(abs(q3d_p - 2.0275884631067768e-16) < 10 ** (-16))
True
```


Tests for q4a (points: 0.5)
```
>>> bool(abs(q4a_t - 5.379366674466881) < 0.0001)
True
```
```
>>> bool(abs(q4a_p - 1.63448584209482e-07) < 10 ** (-7))
True
```


Tests for q4b (points: 0.5)
```
>>> bool(abs(q4b_t + 2.0298173514583273) < 0.0001)
True
```
```
>>> bool(abs(q4b_p - 0.04373374497984475) < 10 ** (-4))
True
```
